# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
import pandas as pd
# df = pd.read_csv('./assets/NISPUF17.csv')
# rate = df["EDUC1"].value_counts()/df["EDUC1"].count()
# s = rate.rename({1:"less than high school", 2:"high school", 3:"more than high school but not college", 4:"college"})    
# new_dict = s.to_dict()
# print(new_dict)
def proportion_of_education():
    df = pd.read_csv('./assets/NISPUF17.csv')
    rate = df["EDUC1"].value_counts()/df["EDUC1"].count()
    s = rate.rename({1:"less than high school", 2:"high school", 3:"more than high school but not college", 4:"college"})    
    new_dict = s.to_dict()
    return new_dict
    raise NotImplementedError()
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
# influenza vaccine recieved / breastmilk 
# influenza vaccine received / no breastmilk
# P_NUMFLU 
# CBF_01
df = pd.read_csv('./assets/NISPUF17.csv')


print((df['CBF_01'].mean(),df['P_NUMFLU'].mean()))
print(round (df['CBF_01'].mean(),1),round(df['P_NUMFLU'].mean(), 1))
def average_influenza_doses():
    df['P_NUMFLU'] = df['P_NUMFLU'].fillna(value=0)
    # number 0 or greater than 0
    
    # remove all breastfeeding status unknown (3,4,99) 
    df['CBF_01'].values[df['CBF_01'].values > 2] = 0
    
    flu_rate = round (df['CBF_01'].mean(), 1)
    breastfed = round(df['P_NUMFLU'].mean(), 1)
    return (flu_rate, breastfed)
    raise NotImplementedError()
```

Output:
```
(1.4681890040400492, 1.8414530750668492)
1.5 1.8
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
# number of kids who contracted chickpox AND were vaccinated against it : vaccinted AND no chicken pox
# return by sex
# "SEX"(1: male, 2:female), "HAD_CPOX" (1-yes, 2-no), "P_NUMVRC" "P_NUMVRN"
df = pd.read_csv('./assets/NISPUF17.csv')

# df['P_NUMVRC']=df['P_NUMVRC'].fillna(value=0)
# df['P_NUMVRN']=df['P_NUMVRN'].fillna(value=0)
# df['HAD_CPOX'].values[df['HAD_CPOX'].values > 2] = 0
# df['SEX'].values[df['SEX'].values > 2] = 0

# PSEUDO
# male_rate = df['SEX'=1]
#     ['HAD_CPOX'=1]
#     ['P_NUMVRN' OR 'P_NUMVRC'>0].count() 
#     / df['SEX'=1]
#     ['HAD_CPOX'=0]
#     ['P_NUMVRN' OR 'P_NUMVRC'>0].count()

# is_vaccinated = df[(df['P_NUMVRN'] >= 1) | (df['P_NUMVRC'] >= 1)]
# male_gotit_rate = len(df[(df["SEX"] == 1) & (df['HAD_CPOX'] ==1) & (df['P_NUMVRC'] >=1)])
# male_nogot_rate = len(df[(df["SEX"] == 1) & (df['HAD_CPOX'] ==0) & (df['P_NUMVRC'] >=1)])
# print(male_gotit_rate, male_nogot_rate)


# male_rate = len(df[
#     (df["SEX"] == 1) & 
#     (df['HAD_CPOX'] ==1) & 
#     (
#         (df['P_NUMVRC'] >=1) |
#         (df['P_NUMVRN'] >=1)
#     )
# ])

# male_rate = male_gotit_rate / male_nogot_rate
# female_rate = df[(df["SEX"] == 2) & (df['HAD_CPOX'] ==1) & (df['P_NUMVRC'] >=1)].count()

# print(rate)


def chickenpox_by_sex():
    df = pd.read_csv('./assets/NISPUF17.csv')
    male_vacc_pos = len(df[
        (df["SEX"] == 1) & 
        (df['HAD_CPOX'] ==1) & 
        (
            (df['P_NUMVRC'] >=1) |
            (df['P_NUMVRN'] >=1)
        )
    ])
    male_vacc_neg = len(df[
        (df["SEX"] == 1) & 
        (df['HAD_CPOX'] == 2) & 
        (
            (df['P_NUMVRC'] >=1) |
            (df['P_NUMVRN'] >=1)
        )
    ])
    print(male_vacc_pos, male_vacc_neg)
    male_rate = male_vacc_pos / male_vacc_neg
    female_vacc_pos = len(df[
        (df["SEX"] == 2) & 
        (df['HAD_CPOX'] ==1) & 
        (
            (df['P_NUMVRC'] >=1) |
            (df['P_NUMVRN'] >=1)
        )
    ])
    female_vacc_neg=  len(df[
        (df["SEX"] == 2) & 
        (df['HAD_CPOX'] ==2) & 
        (
            (df['P_NUMVRC'] >=1) |
            (df['P_NUMVRN'] >=1)
        )
    ])
    print(female_vacc_pos, female_vacc_neg)
    female_rate = female_vacc_pos/female_vacc_neg
    print(male_rate, female_rate)
    return {"male": male_rate,
           "female": female_rate}
    raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

Output:
```
68 7028
53 6802
0.009675583380762664 0.0077918259335489565
```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. If the `had_chickenpox_column` is either `1` (for yes) or `2` for no, and that the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine, then a positive correlation (e.g. `corr > 0`) would mean that an increase in `had_chickenpox_column` (which means more no's) would mean an increase in the `num_chickenpox_vaccine_column` (which means more doses of vaccine). If `corr < 0` then there is a negative correlation, indicating that having had chickenpox is related to an increase in the number of vaccine doses. Also, `pval` refers to the probability the relationship observed is significant. In this case `pval` should be very very small (will end in `e-18` indicating a very small number), which means the result unlikely to be by chance.

[1] This isn't really the full picture, since we are not looking at when the dose was given. It's possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
# import scipy.stats as stats
# import numpy as np
# import pandas as pd

# df = pd.read_csv('./assets/NISPUF17.csv')
# df['P_NUMVRC']=df['P_NUMVRC'].fillna(value=0)
# # df['HAD_CPOX'].values[df['HAD_CPOX'].values > 2] = 0
# corr, pval=stats.pearsonr(df["HAD_CPOX"],df["P_NUMVRC"])
# print(corr, pval)


def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    df = pd.read_csv('./assets/NISPUF17.csv')
    df['P_NUMVRC']=df['P_NUMVRC'].fillna(value=0)

#   here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(df["HAD_CPOX"],df['P_NUMVRC'])
    
    # just return the correlation
    #return corr
    print(corr)
    return corr

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

Output:
```
-0.01026880593978898
```

```python

```